### General Instructions

Do not change the file name, method name or any variable name in your submission file. 

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and student number below.

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Also, ensure that your notebook does not give errors before submitting. Ensure there is no 'Assertion Error', 'NotImplementedError' or  test(s) failed  feedback. 

NotImplementedError: this means there is a code cell/task you are yet to implement.

AssertionError: this means your implementation is failing some tests.

Note that your assignment will be checked with additional test cases after submission. Ensure you work with the instructions given




```python
NAME = "Yaseen Haffejee"
STUDENT_NUMBER = "1827555"
```

---

# <center>  COMS4054A/COMS7066A </center>
# <center> Natural Language Processing/Technology (NLP) 2022 </center>
## <center> Lab Session 3 </center>
### <center> 18 August, 2022 </center>

## N-gram Language Models

#### Total (30 points/marks)

### Objectives
1. To practice the concepts learnt in class
1. To build a simple autocomplete system
1. To use ngram models to generate random text.

```python
#RUN THIS CELL
test1 = """  Sentence 1.\n   Sentence 2.\n Sentence 3.  \n
"""
```

### Task 1 - Sentence Segmentation  (2 points)
1. Split the data using "\n".
2. Remove leading and trailing spaces.
1. Delete sentences if they are emptystrings

Expected Output:
![title](test1.png)

<!-- grade_id: Task1, points: 0, grade: false, solution: true -->
```python
def sentence_segmentation(data):
    """
    perform sentence segmentation using the linebreak "\n" 
    
    Returns: A list of sentences
    """
    
    # Splitting by the newline character
    split_data = data.split("\n")
    # Removing empty strings
    sentences = [x for x in split_data if x!= ""]
    # Removing the leading and trailing spaces
    sentences = [x.strip(" ") for x in sentences]
    return sentences
```

```python
print(sentence_segmentation(test1))
```

Output:
```
['Sentence 1.', 'Sentence 2.', 'Sentence 3.']
```

<!-- grade_id: Task1test, points: 1, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
assert sentence_segmentation(test1) == ['Sentence 1.', 'Sentence 2.', 'Sentence 3.'], "Task 1 test(s) failed!"
print("Test passed")
```

Output:
```
Test passed
```

<!-- grade_id: Task1hiddentest, points: 1, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

### Task 2 - Tokenization (2 points)
1. Perform case folding of text into lowercase
1. Use the nltk.word_tokenize method to split each sentence into a list of tokens
1. Append each set of tokenized sentences into a list of tokenized sentences.

Expected Output:
![title](test2.png)

```python
import nltk
nltk.download('punkt')
```

Output:
```
[nltk_data] Downloading package punkt to
[nltk_data]     C:\Users\yasee\AppData\Roaming\nltk_data...
[nltk_data]   Package punkt is already up-to-date!
True
```

<!-- grade_id: Task2, points: 0, grade: false, solution: true -->
```python
import nltk
def tokenize_sentences(sentences):
    """Args:
    sentences: List of strings
    returns: list of list of tokens"""
    
    # Make everything lowercase
    sentences = [x.lower() for x in sentences]
    # Tokenize the sentences
    tokenized_sentences = [nltk.word_tokenize(x) for x in sentences]
    
    return tokenized_sentences

    
```

<!-- grade_id: Task2test, points: 1, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
assert tokenize_sentences(sentence_segmentation(test1)) == [['sentence', '1', '.'], ['sentence', '2', '.'], ['sentence', '3', '.']], "Task2 test(s) failed!"
print("Tests passed!")
```

Output:
```
Tests passed!
```

<!-- grade_id: Task2hiddentest, points: 1, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

### Task 3 - Count words (2 points)
1. Count all the appearances of a token in the data
2. Create and return a dictionary that maps each token with its count/frequency.

Expected Output:

![title](test3.png)

<!-- grade_id: Task3, points: 0, grade: false, solution: true -->
```python
def count_tokens(tokenized_sentences):
    """Args: list of list of strings (tokens)
    Returns: a dictionary that maps a token (str) to its frequency (int)"""
    # YOUR CODE HERE
    # Create a single list of all the tokens
    all_words = [word for sublist in tokenized_sentences for word in sublist]
    word_counts = {}
    # Get all the tokens and their frequencies
    for word in all_words:
        try:
            word_counts[word] = word_counts[word] + 1
        except:
            word_counts[word] = 1
            
    # Create a dictionary of the token and its frequency
    return word_counts
```

<!-- grade_id: Task3test, points: 1, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
assert count_tokens(tokenize_sentences(sentence_segmentation(test1))) == {'sentence': 3, '1': 1, '.': 3, '2': 1, '3': 1}, "Task 3 test(s) failed!"
print("Test(s) passed!")
```

Output:
```
Test(s) passed!
```

<!-- grade_id: Task3hiddentest, points: 1, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

```python
print( count_tokens(tokenize_sentences(sentence_segmentation(test1))))
```

Output:
```
{'sentence': 3, '1': 1, '.': 3, '2': 1, '3': 1}
```

### Out of vocabulary words

Often, we do not use all the tokens that occur in the data for training, we use frequently occuring words.  Given the words in our training set, we can decide all words that occur a certain number of times would be included in our vocabulary and used in the training. 

So what happens to the other words, a way to handle them is to represent them with an unknown word token.

This way the training data also has some unknown words to train on. If in the test data, we encounter unknown words, we'll use the information we already have about unknown words.



### Task 4 - Vocabulary (2 points)

Write a function such that given a frequency_threshold, returns the list of words that meet the given threshold (equal to  or greater than the threshold).

<!-- grade_id: Task4, points: 0, grade: false, solution: true -->
```python
def get_words_with_frequency(tokenized_sentences, frequency_threshold):
    """Args:
    tokenized sentences: list of list of strings
    frequency_threshold: minimum no of occurrences for a word to be in the vocabulary
    
    returns: a list of words that meet the given threshold."""
    
    vocabulary = []
    # YOUR CODE HERE
    # Get the tokens and their frequencies using the count_tokens method
    tokens_frequencies = count_tokens(tokenized_sentences)
    # Use the threshold to remove any token with count less than the threshold
    for key,value in tokens_frequencies.items():
        if(value>= frequency_threshold):
            vocabulary.append(key)
            
    return vocabulary
    
```

```python
#RUN THIS CELL
tokenized_sentences =[['sentence', '1', '.'], ['sentence', '2', '.'], ['sentence', '3', '.']]

```

<!-- grade_id: Task4_test1, points: 1, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
assert get_words_with_frequency(tokenized_sentences,3) == ['sentence', '.'], "Task 4 test(s) failed!"
print("Test(s) passed!")
```

Output:
```
Test(s) passed!
```

<!-- grade_id: Task4_testhidden, points: 1, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

### Task 5 (2 points)
Given a list of tokenized sentences, replace all words that are not in the given vocabulary with the "\<unk\>" token.



<!-- grade_id: Task5, points: 0, grade: false, solution: true -->
```python
def replace_unknown_words(tokenized_sentences, vocabulary, unknown_token='<unk>'):
    
    """
    Args:
    tokenized_sentences = list of lists of strings
    vocabulary = list of recognized strings
    unknown_token = string to represent out-of-vocabulary words
    
    Returns:
        A list of lists of strings, with words not in the vocabulary replaced
        """
    # YOUR CODE HERE
    replaced_tokenized_sentences = []
    for sentence in tokenized_sentences:
        temp = []
        for word in sentence:
            if word in vocabulary:
                temp.append(word)
            else: 
                temp.append("<unk>")
        replaced_tokenized_sentences.append(temp)

    return replaced_tokenized_sentences
```

<!-- grade_id: Task5test, points: 1, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
vocabulary = ['sentence', '.']
assert replace_unknown_words(tokenized_sentences, vocabulary, unknown_token='<unk>') == [['sentence', '<unk>', '.'], ['sentence', '<unk>', '.'], ['sentence', '<unk>', '.']], "Task 5 test(s) failed!"
print("Test(s) passed!")
```

Output:
```
Test(s) passed!
```

<!-- grade_id: Task5hiddentest, points: 1, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

### Task 6 (2 points)
Calculate all the n-grams that occur in a dataset and the number of times they occur.

The data is presented as a list of list of strings, where each nested list represents a sentence. 

Add 'n' start tokens and 1 end token to each sentence before the calculation.
(Although for a n-gram you need to use n-1 starting markers for a sentence, for this implemtation, prepend 'n' starting markers in order to use them to compute the initial probability for the LM later in the assignment.)

Your function should return a dictionary that maps a tuple of n-words to its frequency

<!-- grade_id: cell-48a08396516d9b47, points: 0, grade: false, solution: true -->
```python
def count_n_grams(data, n, start_token='<s>', end_token = '<e>'):
    """
    Count all n-grams in the data
    
    Args:
        data: List of lists of words
        n: number of words in a sequence
    
    Returns:
        A dictionary that maps a tuple of n-words to its frequency
    """
    
    # Initialize dictionary of n-grams and their counts
    n_grams = {}
    new_data = []
    # Adding the start and end tokens to the sentences
    for sentence in data:
        new_sentence = [start_token]*n
        new_sentence.extend(sentence)
        new_sentence.append(end_token)
        new_data.append(new_sentence)

    # Doing the counting of the n-grams
    for sentence in new_data:
        length = len(sentence)
        for i in range(length-1):
            key = tuple(sentence[i:i+n])
            try:
                n_grams[key] = n_grams[key] + 1

            except:
                n_grams[key] = 1


    return n_grams
```


Expected Output
Bi-gram:

{('\<s\>', '\<s\>'): 2, ('\<s\>', 'i'): 1, ('i', 'like'): 1, ('like', 'green'): 1, ('green', 'apples'): 1, 
 ('apples', '\<e\>'): 1, ('\<s\>', 'this'): 1, ('this', 'apple'): 1, ('apple', 'looks'): 1, ('looks', 'like'): 1, 
 ('like', 'an'): 1, ('an', 'orange'): 1, ('orange', '\<e\>'): 1}

```python
sentences = [['i', 'like', 'green', 'apples'],
             ['this', 'apple', 'looks', 'like', 'an', 'orange']]
print("Uni-gram:")
print(count_n_grams(sentences, 1))
print("Bi-gram:")
print(count_n_grams(sentences, 2))
print("Tri-gram:")
print(count_n_grams(sentences, 3))
```

Output:
```
Uni-gram:
{('<s>',): 2, ('i',): 1, ('like',): 2, ('green',): 1, ('apples',): 1, ('this',): 1, ('apple',): 1, ('looks',): 1, ('an',): 1, ('orange',): 1}
Bi-gram:
{('<s>', '<s>'): 2, ('<s>', 'i'): 1, ('i', 'like'): 1, ('like', 'green'): 1, ('green', 'apples'): 1, ('apples', '<e>'): 1, ('<s>', 'this'): 1, ('this', 'apple'): 1, ('apple', 'looks'): 1, ('looks', 'like'): 1, ('like', 'an'): 1, ('an', 'orange'): 1, ('orange', '<e>'): 1}
Tri-gram:
{('<s>', '<s>', '<s>'): 2, ('<s>', '<s>', 'i'): 1, ('<s>', 'i', 'like'): 1, ('i', 'like', 'green'): 1, ('like', 'green', 'apples'): 1, ('green', 'apples', '<e>'): 1, ('apples', '<e>'): 1, ('<s>', '<s>', 'this'): 1, ('<s>', 'this', 'apple'): 1, ('this', 'apple', 'looks'): 1, ('apple', 'looks', 'like'): 1, ('looks', 'like', 'an'): 1, ('like', 'an', 'orange'): 1, ('an', 'orange', '<e>'): 1, ('orange', '<e>'): 1}
```

```python
sentences = [['I', 'am', 'Sam'],
             ['Sam', 'I', 'am'],
             ["I", "do", "not", "like", "green", 'eggs', 'and', 'ham' ]]
print("Uni-gram:")
print(count_n_grams(sentences, 1))
print("Bi-gram:")
print(count_n_grams(sentences, 2))
print("Tri-gram:")
print(count_n_grams(sentences, 3))
print("Quad-gram:")
print(count_n_grams(sentences, 4))
print("Quint-gram:")
print(count_n_grams(sentences, 5))
```

Output:
```
Uni-gram:
{('<s>',): 3, ('I',): 3, ('am',): 2, ('Sam',): 2, ('do',): 1, ('not',): 1, ('like',): 1, ('green',): 1, ('eggs',): 1, ('and',): 1, ('ham',): 1}
Bi-gram:
{('<s>', '<s>'): 3, ('<s>', 'I'): 2, ('I', 'am'): 2, ('am', 'Sam'): 1, ('Sam', '<e>'): 1, ('<s>', 'Sam'): 1, ('Sam', 'I'): 1, ('am', '<e>'): 1, ('I', 'do'): 1, ('do', 'not'): 1, ('not', 'like'): 1, ('like', 'green'): 1, ('green', 'eggs'): 1, ('eggs', 'and'): 1, ('and', 'ham'): 1, ('ham', '<e>'): 1}
Tri-gram:
{('<s>', '<s>', '<s>'): 3, ('<s>', '<s>', 'I'): 2, ('<s>', 'I', 'am'): 1, ('I', 'am', 'Sam'): 1, ('am', 'Sam', '<e>'): 1, ('Sam', '<e>'): 1, ('<s>', '<s>', 'Sam'): 1, ('<s>', 'Sam', 'I'): 1, ('Sam', 'I', 'am'): 1, ('I', 'am', '<e>'): 1, ('am', '<e>'): 1, ('<s>', 'I', 'do'): 1, ('I', 'do', 'not'): 1, ('do', 'not', 'like'): 1, ('not', 'like', 'green'): 1, ('like', 'green', 'eggs'): 1, ('green', 'eggs', 'and'): 1, ('eggs', 'and', 'ham'): 1, ('and', 'ham', '<e>'): 1, ('ham', '<e>'): 1}
Quad-gram:
{('<s>', '<s>', '<s>', '<s>')
```

<!-- grade_id: Task6test, points: 1, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
sentences = [['I', 'am', 'Sam'],
             ['Sam', 'I', 'am'],
             ["I", "do", "not", "like", "green", 'eggs', 'and', 'ham' ]]

assert count_n_grams(sentences, 2) == {('<s>', '<s>'): 3, ('<s>', 'I'): 2, ('I', 'am'): 2, ('am', 'Sam'): 1, ('Sam', '<e>'): 1, ('<s>', 'Sam'): 1, ('Sam', 'I'): 1, ('am', '<e>'): 1, ('I', 'do'): 1, ('do', 'not'): 1, ('not', 'like'): 1, ('like', 'green'): 1, ('green', 'eggs'): 1, ('eggs', 'and'): 1, ('and', 'ham'): 1, ('ham', '<e>'): 1}, "Task 6 test(s) failed!"
print("Test(s) passed!")
```

Output:
```
Test(s) passed!
```

<!-- grade_id: Task6hiddentest, points: 1, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

### Task 7 (2 points)
Let's estimate the n-gram probabilities using the formula:

![title](image.png)

Write a function that returns the n-gram probabilities, given the n-gram counts.

<!-- grade_id: Task7, points: 0, grade: false, solution: true -->
```python
def estimate_probability_raw(word, n_minus_1_gram, 
                         n_minus_1_gram_counts, n_gram_counts, vocabulary_size):
    """
    Estimate the probabilities of a next word using the n-gram counts with k-smoothing
    
    Args:
        word: next word (estimating its probability)
        n_minus_1_gram: A sequence of words of length n-1
        n_minus_1_gram_counts: Dictionary of counts of (n-1)-grams
        n_gram_counts: Dictionary of counts of n-grams
        vocabulary_size: number of words in the vocabulary
    
    Returns:
        A probability
    """
    # convert list to tuple to use it as a dictionary key
    n_minus_1_gram = tuple(n_minus_1_gram)
    
    numerator = n_gram_counts[n_minus_1_gram+tuple([word])]
    denominator = n_minus_1_gram_counts[n_minus_1_gram]
    probability = numerator/denominator
    return probability
```

<!-- grade_id: Task7_test, points: 1, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
# Expected output = 0.33

sentences = [['I', 'am', 'Sam'],
             ['Sam', 'I', 'am'],
             ["I", "do", "not", "like", "green", 'eggs', 'and', 'ham' ]]
unique_words = list(set(sentences[0] + sentences[1]+sentences[2]))
previous = ['am']
unigram_counts = count_n_grams(sentences, 1)

bigram_counts = count_n_grams(sentences, 2)
prob = estimate_probability_raw("do", ['I'], unigram_counts, bigram_counts, len(unique_words))

assert round(prob,2) == 0.33, "Task 7 test(s) failed!"
print("Test(s) passed!")
```

Output:
```
Test(s) passed!
```

<!-- grade_id: Task7hiddentest, points: 1, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

### Task 8 - Laplace Smoothing (2 points)
Create a function that returns the probabilities of the n-grams after perfroming add-1 smoothing using the formula:

![title](image_2.png)
    

<!-- grade_id: Task8test, points: 0, grade: false, solution: true -->
```python
def estimate_probability(word, n_minus_1_gram, 
                         n_minus_1_gram_counts, n_gram_counts, vocabulary_size, k=1.0):
    """
    Estimate the probabilities of a next word using the n-gram counts with k-smoothing
    
    Args:
        word: next word
        n_minus_1_gram: A sequence of words of length n-1
        n_minus_1_gram_counts: Dictionary of counts of (n-1)grams
        n_gram_counts: Dictionary of counts of n-grams
        vocabulary_size: number of words in the vocabulary
        k: positive constant, smoothing parameter
    
    Returns:
        A probability
    """
    # convert list to tuple to use it as a dictionary key
    n_minus_1_gram = tuple(n_minus_1_gram)
    # We use the try and except in case an n-gram has never been seen before
    try: 
        numerator = n_gram_counts[n_minus_1_gram+tuple([word])]
    except:
        numerator = 0
    try:
        denominator = n_minus_1_gram_counts[n_minus_1_gram]
    except: 
        denominator = 0
    probability = (numerator+k)/(denominator+vocabulary_size)
    
    return probability
```

<!-- grade_id: Task8_test, points: 1, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
# Expected output = 0.15

sentences = [['I', 'am', 'Sam'],
             ['Sam', 'I', 'am'],
             ["I", "do", "not", "like", "green", 'eggs', 'and', 'ham' ]]
unique_words = list(set(sentences[0] + sentences[1]+sentences[2]))
unigram_counts = count_n_grams(sentences, 1)

bigram_counts = count_n_grams(sentences, 2)
prob_laplace = estimate_probability("do", ['I'], unigram_counts, bigram_counts, len(unique_words))
assert round(prob_laplace,2) == 0.15, "Task 8 test(s) failed!"
print("Test(s) passed!")
```

Output:
```
Test(s) passed!
```

<!-- grade_id: Task8_hiddentest, points: 1, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

### Estimate probabilities for all words (1 point)
The function defined below loops over all words in vocabulary to calculate probabilities for all possible words.

This function is provided for you.

The expected output after running the test is:

![title](image4.png)

```python
def estimate_probabilities(n_minus_1_gram, n_minus_1_gram_counts, n_gram_counts, vocabulary, end_token='<e>', unknown_token="<unk>",  k=1.0):
    """
    Estimate the probabilities of next words using the n-gram counts with k-smoothing
    
    Args:
        n_minus_1_gram: A sequence of words of length n-1
        n_minus_1_gram_counts: Dictionary of counts of (n-1)grams
        n_gram_counts : Dictionary of counts of (n)-grams
        vocabulary: List of words
        k: positive constant, smoothing parameter
    
    Returns:
        A dictionary mapping from next words to the probability.
    """
    # convert list to tuple to use it as a dictionary key
    n_minus_1_gram = tuple(n_minus_1_gram)    
    
    # add <e> <unk> to the vocabulary
        # <s> is not needed since it should not appear as the next word
    vocabulary = vocabulary + [end_token, unknown_token]    
    vocabulary_size = len(vocabulary)    
    
    probabilities = {}
    for word in vocabulary:
        probability = estimate_probability(word, n_minus_1_gram, 
                                           n_minus_1_gram_counts, n_gram_counts , 
                                           vocabulary_size, k=k)
                
        probabilities[word] = probability

    return probabilities
```

<!-- grade_id: cell-83e83295465c517a, points: 1, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
sentences = [['i', 'like', 'a', 'cat'],
             ['this', 'dog', 'is', 'like', 'a', 'cat']]
unique_words = list(set(sentences[0] + sentences[1]))
unigram_counts = count_n_grams(sentences, 1)
bigram_counts = count_n_grams(sentences, 2)

assert estimate_probabilities("a", unigram_counts, bigram_counts, unique_words, k=1) == {'i': 0.09090909090909091,
 'is': 0.09090909090909091,
 'cat': 0.2727272727272727,
 'this': 0.09090909090909091,
 'a': 0.09090909090909091,
 'like': 0.09090909090909091,
 'dog': 0.09090909090909091,
 '<e>': 0.09090909090909091,
 '<unk>': 0.09090909090909091}, "Task 8 Test failed"
print("Test passed")
```

Output:
```
Test passed
```

### Task 9 (4 points)
Write a function to compute the probabilites of all possible next words and return the word with the highest probability, alongside its probability as a tuple.


<!-- grade_id: cell-4cae473dd932f2b1, points: 0, grade: false, solution: true -->
```python
def suggest_a_word(previous_tokens,n_minus_1_gram_counts , n_gram_counts , vocabulary, end_token='<e>', unknown_token="<unk>", k=1.0, start_with=None):
    """
    Get suggestion for the next word
    
    Args:
        previous_tokens: The sentence you input where each token is a word. Must have length >= n -1
        n_minus_1_gram_counts: Dictionary of counts of (n-1)-grams
        n_gram_counts: Dictionary of counts of (n)-grams
        vocabulary: List of words
        k: positive constant, smoothing parameter
        start_with: If not None, specifies the first few letters of the next word
        
    Returns:
        A tuple of 
          - string of the most likely next word
          - corresponding probability
    """
    # YOUR CODE HERE
    ## Get n which is how many words we are considering for the history
    n = len(list(n_gram_counts.keys())[0])
    n_minus_one = n-1
    ## From the incoming words that we want to predict the next word for, we can the n-1 words
    n_minus_1_gram = previous_tokens[-n_minus_one:]
    ## caluclate the probability of each word in the vocab being the next word
    probabilities = estimate_probabilities(n_minus_1_gram, n_minus_1_gram_counts, n_gram_counts, vocabulary, k=k)

    if (start_with != None):
        items_starting_with = [ (k,v) for k,v in probabilities.items() if k.startswith(start_with)]
        keys,vals = zip(*items_starting_with)
        idx = vals.index(max(vals))
        suggestion = keys[idx]
        max_prob = vals[idx]

    else:
        suggestion = max(probabilities, key=probabilities.get)
        max_prob = probabilities[suggestion]
       
    return suggestion, max_prob
```

<!-- grade_id: Task9test, points: 2, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
sentences = [['i', 'like', 'a', 'cat'],
             ['this', 'dog', 'is', 'like', 'a', 'cat']]
unique_words = list(set(sentences[0] + sentences[1]))

unigram_counts = count_n_grams(sentences, 1)
bigram_counts = count_n_grams(sentences, 2)

previous_tokens = ["i", "like"]


# test your code when setting the starts_with
tmp_starts_with = 'c'
tmp_suggest2 = suggest_a_word(previous_tokens, unigram_counts, bigram_counts, unique_words, k=1.0, start_with=tmp_starts_with)


assert tmp_suggest2 == ('cat', 0.09090909090909091), "Task 9 Test failed, Be sure to implement using startswith"
print("Test passed")
```

Output:
```
Test passed
```

<!-- grade_id: cell-7854fc39a5cf1747, points: 2, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

### Additional Task 1 (2 points)

By the suggest_a_word function, we have a simple autocomplete system, such that given previous tokens, it can predict the next word.
Although we are predicting the word with the highest frequency, we can modify the function to suggest any random word.

Write a function that suggests a random word and not the word with the highest probability.



```python
import random
import numpy as np

def suggest_a_random_word(previous_tokens,n_minus_1_gram_counts , n_gram_counts , vocabulary, end_token='<e>', unknown_token="<unk>", k=1.0, start_with=None):
    """
    Get suggestion for the next word randomly
    
    Args:
        previous_tokens: The sentence you input where each token is a word. Must have length >= n -1
        n_minus_1_gram_counts: Dictionary of counts of (n-1)-grams
        n_gram_counts: Dictionary of counts of (n)-grams
        vocabulary: List of words
        k: positive constant, smoothing parameter
        start_with: If not None, specifies the first few letters of the next word
        
    Returns:
        A tuple of 
          - string of a random word
          - corresponding probability
    """
    # YOUR CODE HERE
    ## Get n which is how many words we are considering for the history
    n = len(list(n_gram_counts.keys())[0])
    n_minus_one = n-1
    ## From the incoming words that we want to predict the next word for, we can the n-1 words
    n_minus_1_gram = previous_tokens[-n_minus_one:]
    ## caluclate the probability of each word in the vocab being the next word
    probabilities = estimate_probabilities(n_minus_1_gram, n_minus_1_gram_counts, n_gram_counts, vocabulary, k=k)

    if(start_with == None):
        # word,prob = random.choice(list(probabilities.items()))
        word = np.random.choice(a = list(probabilities.keys()))
        prob = probabilities[word]
    else:
        items_starting_with = [ (k,v) for k,v in probabilities.items() if k.startswith(start_with)]
        word,prob = random.choice(items_starting_with)

    return word,prob

```

```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
sentences = [['i', 'like', 'a', 'cat'],
             ['this', 'dog', 'is', 'like', 'a', 'cat']]
unique_words = list(set(sentences[0] + sentences[1]))

unigram_counts = count_n_grams(sentences, 1)
bigram_counts = count_n_grams(sentences, 2)

previous_tokens = ["i", "like"]


# test your code when setting the starts_with
tmp_starts_with = None
tmp_suggest2 = suggest_a_random_word(previous_tokens, unigram_counts, bigram_counts, unique_words, k=1.0, start_with=tmp_starts_with)

print(f"The randomly suggested word is: {tmp_suggest2[0]} and the corresponding probability is : {tmp_suggest2[1]}")
```

Output:
```
The randomly suggested word is: is and the corresponding probability is : 0.09090909090909091
```

### Additional Task 2 (7 points)
You can also put all the functions we have written together to autogenerate random texts of different lengths using different n-grams.
You can also download the sample text of e-books here: https://www.gutenberg.org/ebooks/
        
Write a function that generates random texts of different lengths, using different n-grams.
1. The function should take in a parameter that specifies the length of words to be generated 
1. The function should take in a parameter that specifies the n-gram to be used (i.e. the function should work using different ngrams).
1. Use a body of raw text not less than 5000 words
1. Preprocess the raw text using all the previous functions in the previous tasks.
1. Write a function to calculate the perplexity metric using the function definition given below. 

![title](perplexity.png)

# Corpus

### Pre-processing the corpus

```python
with open("tmp.txt", "r",encoding="UTF-8") as txt_file:
 story = txt_file.readlines()

 data = [sentence_segmentation(s) for s in story if s!="\n"]
 tokenized_sentences = [tokenize_sentences(s) for s in data if "Chapter" not in s[0]]

sentences = []
for t in tokenized_sentences:
    try:
        ## Sentences will be a list of lists where each sublist is a sentence
        sentences.extend(t)
    except:
        continue
```

```python
frequencies = count_tokens(sentences)
total_number_of_words_in_corpus = sum(frequencies.values())

assert (total_number_of_words_in_corpus >= 5000),"Ensure Corpus has more than 5000 words"
print("Test Passed !")
print(f"Corpus has {total_number_of_words_in_corpus} words.")
```

Output:
```
Test Passed !
Corpus has 47044 words.
```

```python
story_as_sentences = sentences

## Getting the vocabulary of the corpus. A words is only part of the vocabulary if it occurs 3 or more times
vocabulary = get_words_with_frequency(story_as_sentences,3)

## Replacing words that are not in the vocabulary with <unk> tag
story_as_sentences = replace_unknown_words(story_as_sentences,vocabulary)
```

```python
def generate_random_texts(length_of_words_to_be_generated,n,data,vocabulary):
    """
    Generate a random sentence
    
    Args:
        length_of_words_to_be_generated: integer denoting the number of words to be generated
        n: integer denoting which n-gram to use
        data: A list of lists where ach sublist contains a sentences of the corpus
        vocabulary: The unique words of the corpus
        k: Positive smoothing constant
    
    Returns:
       length_of_words_to_be_generated randomly generated sentence
    """
    sentences = []
    ## Create the n_gram and n-1 gram
    n_gram_counts = count_n_grams(data,n)
    n_minus_one_gram_counts = count_n_grams(data,n-1)
    i = 0
    unkown= "<unk>"
    stop = "<e>"
    start = "<s>"
    found = True
    while(found):
     ## randomly choose an n-1 gram
        n_minus_one_gram = list(random.choice(list(n_minus_one_gram_counts.keys())))
        # if(unkown in n_minus_one_gram or start in n_minus_one_gram or stop in n_minus_one_gram):
        if(unkown in n_minus_one_gram):

            continue
        else:
            found = False
    ## Loop through and create length_of_words_to_be_generated number of words
    while(i < length_of_words_to_be_generated):
        ## Based on the n-1 gram, suggest a random word from the vocabulary
        suggestion = suggest_a_random_word(n_minus_one_gram,n_minus_one_gram_counts,n_gram_counts,vocabulary)[0]
        ## If the sugested word is any of the tags, make another suggestion
        if(suggestion == unkown or suggestion == start or suggestion ==  stop):
            continue
        i += 1
        ## store the generated sentence
        sentences.append(suggestion)
        ## Add the generated word to the n-1 gram
        # Re-assign the n-1 gram
        n_minus_one_gram.pop(0)
        n_minus_one_gram.append(suggestion)
    return " ".join(sentences)

```

```python
n = 6
length = 4
random_text = generate_random_texts(length,n,story_as_sentences,vocabulary)
print(f"We asked for {length} randomly generated words using {n}-gram.")
print(f"The sentence is: {random_text}")
```

Output:
```
We asked for 4 randomly generated words using 6-gram.
The sentence is: bridges kalidahs loving desire
```

```python
def calculate_perplexity(sentence, n_minus_1_gram_counts, n_gram_counts, vocabulary_size, start_token, end_token, k=1.0):
    """
    Calculate perplexity for a list of sentences
    
    Args:
        sentence: List of strings
        n_minus_1_gram_counts: Dictionary of counts of (n-1)-grams
        n_gram_counts: Dictionary of counts of (n)-grams
        vocabulary_size: number of unique words in the vocabulary
        k: Positive smoothing constant
    
    Returns:
        Perplexity score to 4 decimal places
    """
    
    
    #YOUR CODE HERE
    perplexity = 0.0000
    n = len(list(n_minus_1_gram_counts.keys())[0])
    sentence = [start_token]*n + sentence + [end_token]*n
    sentence = tuple(sentence)
    N = len(sentence)
    total_probability = 1.0
    for i in range(n,N):
        n_minus_one_gram = sentence[i-n:i]
        word = sentence[i]
        probs = estimate_probability(word,n_minus_one_gram,n_minus_1_gram_counts,n_gram_counts,vocabulary_size,k)
        total_probability *= 1/probs

    perplexity = np.round(total_probability**(1/N),4)
    return perplexity
```

```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
unigram_counts = count_n_grams(story_as_sentences, 1)
bigram_counts = count_n_grams(story_as_sentences, 2)

text_from_book = "When Dorothy stood in the doorway and looked around"
perplexity1 = calculate_perplexity(random_text.split(" "), unigram_counts, bigram_counts, len(vocabulary), start_token = "<s>", end_token = "<e>", k=1.0)
perplexity2 = calculate_perplexity(text_from_book.split(" "), unigram_counts, bigram_counts, len(vocabulary), start_token = "<s>", end_token = "<e>", k=1.0)
print(f"The perplexity for the randomly generated sentence: {random_text}, is: {perplexity1}")
print(f"The perplexity for a sentence from the book: {text_from_book}, is: {perplexity2}")
```

Output:
```
The perplexity for the randomly generated sentence: bridges kalidahs loving desire, is: 479.1405
The perplexity for a sentence from the book: When Dorothy stood in the doorway and looked around, is: 291.2747
```

```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
sentences = [['i', 'like', 'a', 'cat'],
             ['this', 'dog', 'is', 'like', 'a', 'cat']]
unique_words = list(set(sentences[0] + sentences[1]))

unigram_counts = count_n_grams(sentences, 1)
bigram_counts = count_n_grams(sentences, 2)


perplexity1 = calculate_perplexity(sentences[0], unigram_counts, bigram_counts, len(unique_words), start_token = "<s>", end_token = "<e>", k=1.0)
perplexity2 = calculate_perplexity(sentences[1], unigram_counts, bigram_counts, len(unique_words), start_token = "<s>", end_token = "<e>", k=1.0)
unseen = ["i","dog","like"]
unseen_perplexity = calculate_perplexity(unseen, unigram_counts, bigram_counts, len(unique_words), start_token = "<s>", end_token = "<e>", k=1.0)
print(f"The perplexity for sentence 1 is: {perplexity1}")
print(f"The perplexity for sentence 2 is: {perplexity2}")
print(f"The perplexity for the unseen sentence is: {unseen_perplexity}")

```

Output:
```
The perplexity for sentence 1 is: 2.804
The perplexity for sentence 2 is: 3.0644
The perplexity for the unseen sentence is: 4.8164
```