### General Instructions

Do not change the file name, method name or any variable name in your submission file. 

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and student number below.

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Also, ensure that your notebook does not give errors before submitting. Ensure there is no 'Assertion Error', 'NotImplementedError' or  test(s) failed  feedback. 

NotImplementedError: this means there is a code cell/task you are yet to implement.

AssertionError: this means your implementation is failing some tests.

Note that your assignment will be checked with additional test cases after submission. Ensure you work with the instructions given




```python
NAME = "Yaseen Haffejee"
STUDENT_NUMBER = "1827555"
```

---

# <center>  COMS4054A/COMS7066A </center>
# <center> Natural Language Processing/Technology (NLP) 2022 </center>
## <center> Lab Session 2 </center>
### <center> 4 August, 2022 </center>

## Text Normalization and Edit Distance

### Objectives

The goal of this session is to have a practical engagement on the topics discussed in Lecture 2.

- Perform different text normalization tasks using raw Python code as well as the NLTK library.
   - Removal of punctuations
   - Tokenization
   - Stemming
   - Lemmatization etc.
- Implement the minimum edit distance algorithm.
- Build a spell checker using the minimum edit distance algorithm.

### Outline
- Lesson 1 - Text Normalization
- Lesson 2 - Minimum Edit Distance and Spell Checker

### Instructions

- Think about it and implement your method/approach before looking at the hints.
- The hints are just one way to look at the tasks.
- You do not have to follow the approach given in the hint.
- Download the required files (testing.txt, correct_words.txt, incorrect_words.txt)

### Marks Allocation (Total - 30 points)
Tasks 1, 2 and 3: 10 points

Task 4: 6 points

Task 5: 5 points

Task 7: 9 points

### Lesson 1 - Text Normalization

#### Question 1.1

Write a simple data preprocessing pipeline in raw Python code. Do not use any libraries.
Given a piece of text in English language, your program should:
- Remove punctuations (Use string.punctuation in python for the list of all punctuations)
- Tokenize the string of text (split by spaces)
- Perform casefolding (turn all text to lowercase)
- Return the cleaned tokens as a list
    



```python
with open('testing.txt', 'r') as file:
    data = file.read().replace('\n'," ")
data
```

Output:
```
"You! I came here yesterday, but I didn't see you. So where were you? She is eighty-seven years today. I am learning NLP, I love love NLP"
```

#### Task 1: Write code to remove punctuation

<details>
<summary>
    <font size="3" color="darkgreen"><b>Hints</b></font>
</summary>
<p>
You can iterate through the text character by character, then join together afterwards    

</p>
</details>

Expected Output:

![title](img1.png)

<!-- grade_id: Task1, points: 0, grade: false, solution: true -->
```python
import string
test_string = data

# this function should return your text with punctuations removed

def task1(test_string):
    # YOUR CODE HERE
    punct_removed = test_string.translate(test_string.maketrans("","",string.punctuation))
    return punct_removed
    raise NotImplementedError()
    
punct_removed = task1(test_string)
```

<!-- grade_id: Task1_1, points: 1, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
assert task1(test_string) == 'You I came here yesterday but I didnt see you So where were you She is eightyseven years today I am learning NLP I love love NLP',"Task 1 test(s) failed"
print("Test(s) passed")
```

Output:
```
Test(s) passed
```

<!-- grade_id: Task1_2, points: 1, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

#### Task 2: Write code to remove the spaces

Write it as a function. 

Input = output from task 1

Output = test string without spaces, returned as a list
<details>
<summary>
    <font size="3" color="darkgreen"><b>Hints</b></font>
</summary>
<p>
string has a split() method.

</p>
</details>

Expected Output:

![title](img2.png)

<!-- grade_id: Task2, points: 0, grade: false, solution: true -->
```python
def task2(punct_removed):
    # YOUR CODE HERE
    spaces_removed = punct_removed.split(" ")
    return spaces_removed
    raise NotImplementedError()
spaces_removed = task2(punct_removed)
print(spaces_removed)
```

Output:
```
['You', 'I', 'came', 'here', 'yesterday', 'but', 'I', 'didnt', 'see', 'you', 'So', 'where', 'were', 'you', 'She', 'is', 'eightyseven', 'years', 'today', 'I', 'am', 'learning', 'NLP', 'I', 'love', 'love', 'NLP']
```

<!-- grade_id: Task2_1, points: 1, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
assert task2(task1(test_string)) == ['You', 'I', 'came', 'here', 'yesterday', 'but', 'I', 'didnt', 'see', 'you', 'So', 'where', 'were', 'you', 'She', 'is', 'eightyseven', 'years', 'today', 'I', 'am', 'learning', 'NLP', 'I', 'love', 'love', 'NLP'] ,"Task 2 test(s) failed"
print("Test(s) passed")
```

Output:
```
Test(s) passed
```

#### Task 3: Write code to perform case folding



Expected Output:

![title](img3.png)

<!-- grade_id: Task3, points: 0, grade: false, solution: true -->
```python
def task3(spaces_removed):
    to_lower_case = []

    for word in spaces_removed:
        to_lower_case.append(word.lower())#type code here

    return to_lower_case # this variable should contain your output
    raise NotImplementedError()
```

<!-- grade_id: Task3_1, points: 3, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
assert task3(task2(task1(test_string))) == ['you', 'i', 'came', 'here', 'yesterday', 'but', 'i', 'didnt', 'see', 'you', 'so', 'where', 'were', 'you', 'she', 'is', 'eightyseven', 'years', 'today', 'i', 'am', 'learning', 'nlp', 'i', 'love', 'love', 'nlp'],"Task 3 test(s) failed"
print("Test(s) passed")
```

Output:
```
Test(s) passed
```

<!-- grade_id: Task1_2_3, points: 4, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

#### Question 1.2
- Use the NLTK to perform the same task
- Ensure you get exactly the same results.

```python
import nltk
# nltk.download("punkt")

from nltk.tokenize import word_tokenize
token_punkt = word_tokenize(test_string, preserve_line=True)
print(token_punkt)
```

Output:
```
['You', '!', 'I', 'came', 'here', 'yesterday', ',', 'but', 'I', 'did', "n't", 'see', 'you.', 'So', 'where', 'were', 'you', '?', 'She', 'is', 'eighty-seven', 'years', 'today.', 'I', 'am', 'learning', 'NLP', ',', 'I', 'love', 'love', 'NLP']
```

How does this work? 
It separated the punctuation marks as tokens.
How do we remove the punctuation marks? 
There are different ways around this with different implications.

```python
def remove_punct(token):
    return [word for word in token if word.isalpha()]

token_punkt_ = remove_punct(token_punkt)
print(token_punkt_)
```

Output:
```
['You', 'I', 'came', 'here', 'yesterday', 'but', 'I', 'did', 'see', 'So', 'where', 'were', 'you', 'She', 'is', 'years', 'I', 'am', 'learning', 'NLP', 'I', 'love', 'love', 'NLP']
```


Something is missing. Is this the output we want?

It is good to know that differen tokenizers behave differently. Always test and use the one that suits your purpose. 


```python
#Different tokenizers behave differently
# see http://text-processing.com/demo/tokenize/

from nltk.tokenize import RegexpTokenizer
tokenizer = RegexpTokenizer(r'\w+')
tokenizer.tokenize(test_string)
```

Output:
```
['You',
 'I',
 'came',
 'here',
 'yesterday',
 'but',
 'I',
 'didn',
 't',
 'see',
 'you',
 'So',
 'where',
 'were',
 'you',
 'She',
 'is',
 'eighty',
 'seven',
 'years',
 'today',
 'I',
 'am',
 'learning',
 'NLP',
 'I',
 'love',
 'love',
 'NLP']
```

#### Task 4: Modify the output.

Let's assume we want exactly the same behaviour as was implemented in Question 1.

Write a function that takes in the output from the TweekTokenizer, 'tweet_tokenizer_tokens', processes it (remove punctuations etc.) and gives exactly the same output as when we tokenized manually. In other words, your output should be the same as 'to_lower_case' from Question 1.1.

Expected Output:
![title](img4.png)


<!-- grade_id: Task4, points: 0, grade: false, solution: true -->
```python
from nltk.tokenize import TweetTokenizer
import re
tweet_tokenizer = TweetTokenizer()
tweet_tokenizer_tokens = tweet_tokenizer.tokenize(test_string)
print(tweet_tokenizer_tokens)

def remove_punct2(token):
    list_words = []
    # YOUR CODE HERE
    for word in token:
        ## Use regex to remove any special characters which denotes the punctuations
        w = re.sub("\W","",word)
        if(w != ""):
            list_words.append(w.lower())
    return list_words
    raise NotImplementedError()


tokens_twitter = remove_punct2(tweet_tokenizer_tokens)
print(tokens_twitter)
```

Output:
```
['You', '!', 'I', 'came', 'here', 'yesterday', ',', 'but', 'I', "didn't", 'see', 'you', '.', 'So', 'where', 'were', 'you', '?', 'She', 'is', 'eighty-seven', 'years', 'today', '.', 'I', 'am', 'learning', 'NLP', ',', 'I', 'love', 'love', 'NLP']
['you', 'i', 'came', 'here', 'yesterday', 'but', 'i', 'didnt', 'see', 'you', 'so', 'where', 'were', 'you', 'she', 'is', 'eightyseven', 'years', 'today', 'i', 'am', 'learning', 'nlp', 'i', 'love', 'love', 'nlp']
```

<!-- grade_id: Task4_hidden, points: 3, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS

```

<!-- grade_id: Task4_1, points: 3, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
assert remove_punct2(tweet_tokenizer_tokens) == ['you', 'i', 'came', 'here', 'yesterday', 'but', 'i', 'didnt', 'see', 'you', 'so', 'where', 'were', 'you', 'she', 'is', 'eightyseven', 'years', 'today', 'i', 'am', 'learning', 'nlp', 'i', 'love', 'love', 'nlp'],"Task 4 test(s) failed"
print("Test(s) passed")
```

Output:
```
Test(s) passed
```

```python
#test both strings and report if they are the same or not.
print(tokens_twitter == to_lower_case)
```

Output:
```
True
```

#### Question 1.3 - Stemming and Lemmatization
Let's use the PorterStemmer and SnowballStemmer on a sample tokenized text.
You can also try them out on your final output from the previous question. 

```python
tweet_tokenizer = TweetTokenizer()
test_string_2 = "Tokenization is the process of segmenting running text into sentences and words. In essence, it's the task of cutting a text into pieces called tokens."
test_string_2_tokens = tweet_tokenizer.tokenize(test_string_2)
print(test_string_2_tokens)

test_string_2_tokens = remove_punct2(test_string_2_tokens)
print(test_string_2_tokens)
```

Output:
```
['Tokenization', 'is', 'the', 'process', 'of', 'segmenting', 'running', 'text', 'into', 'sentences', 'and', 'words', '.', 'In', 'essence', ',', "it's", 'the', 'task', 'of', 'cutting', 'a', 'text', 'into', 'pieces', 'called', 'tokens', '.']
['tokenization', 'is', 'the', 'process', 'of', 'segmenting', 'running', 'text', 'into', 'sentences', 'and', 'words', 'in', 'essence', 'its', 'the', 'task', 'of', 'cutting', 'a', 'text', 'into', 'pieces', 'called', 'tokens']
```

```python
from nltk.stem import PorterStemmer
ps = PorterStemmer()
ps_stem_text = [ps.stem(words_sent) for words_sent in test_string_2_tokens]
print(ps_stem_text)
```

Output:
```
['token', 'is', 'the', 'process', 'of', 'segment', 'run', 'text', 'into', 'sentenc', 'and', 'word', 'in', 'essenc', 'it', 'the', 'task', 'of', 'cut', 'a', 'text', 'into', 'piec', 'call', 'token']
```

There's also the Snowball Stemmer

```python
from nltk.stem import SnowballStemmer
sb = SnowballStemmer("english")
sb_stem_sent = [sb.stem(words_sent) for words_sent in test_string_2_tokens]
print(sb_stem_sent)
```

Output:
```
['token', 'is', 'the', 'process', 'of', 'segment', 'run', 'text', 'into', 'sentenc', 'and', 'word', 'in', 'essenc', 'it', 'the', 'task', 'of', 'cut', 'a', 'text', 'into', 'piec', 'call', 'token']
```

### Notes for myself
Stemming allows us to reduce words into their root form. This ensures we do not have duplicates in our data. E.g, connects, connecting have the same root word of connect.<br>
Sometimes stemming may produce us with words/stems that are not actually part of the language

#### Question 1.4 - Lemmatization

Let's use the WordNetLemmatizer.

```python
nltk.download ("wordnet")
nltk.download('omw-1.4')
```

Output:
```
[nltk_data] Downloading package wordnet to
[nltk_data]     C:\Users\yasee\AppData\Roaming\nltk_data...
[nltk_data]   Package wordnet is already up-to-date!
[nltk_data] Downloading package omw-1.4 to
[nltk_data]     C:\Users\yasee\AppData\Roaming\nltk_data...
[nltk_data]   Package omw-1.4 is already up-to-date!
True
```

```python
from nltk.stem.wordnet import WordNetLemmatizer
lemmatizer = WordNetLemmatizer()
lem_sent = [lemmatizer.lemmatize(words_sent) for words_sent in test_string_2_tokens]
print(lem_sent)
```

Output:
```
['tokenization', 'is', 'the', 'process', 'of', 'segmenting', 'running', 'text', 'into', 'sentence', 'and', 'word', 'in', 'essence', 'it', 'the', 'task', 'of', 'cutting', 'a', 'text', 'into', 'piece', 'called', 'token']
```

#### Just some notes...

Did you notice the differences between stemming and lemmatization. 
Try out other words or sentences.

- Stemming simply chops off the end of the word, based on some heuristics
- The stem word does not have to be a real word in the given language. E.g. datum can be reduced to datu (which is not in the English dictionary).
- Lemmatization reduces to the base word, it ensures the root word belongs to the given language
- Lemmatization hurts performance though

Is text normalization always required or beneficial?

- Yes, text normalisation is always required since it allows us to reduce the randomness in our data.By reducing the randomness, we essentially have a standard for the text which the computer will process. Having this standard will allow us to have more efficient models.

#### Question 1.5 - Removing Stop words

Sometimes, you might also need to remove stop words (common words) from the text.

The NLTK has a set of stopwords, let's download that and remove it from a sample text.


```python
import nltk # Natural Language tool kit 
nltk.download('stopwords')

# You have to download stopwords Package to execute this command
from nltk.corpus import stopwords

# check out the stop words in English
stopwords.words('english')
```

Output:
```
[nltk_data] Downloading package stopwords to
[nltk_data]     C:\Users\yasee\AppData\Roaming\nltk_data...
[nltk_data]   Package stopwords is already up-to-date!
['i',
 'me',
 'my',
 'myself',
 'we',
 'our',
 'ours',
 'ourselves',
 'you',
 "you're",
 "you've",
 "you'll",
 "you'd",
 'your',
 'yours',
 'yourself',
 'yourselves',
 'he',
 'him',
 'his',
 'himself',
 'she',
 "she's",
 'her',
 'hers',
 'herself',
 'it',
 "it's",
 'its',
 'itself',
 'they',
 'them',
 'their',
 'theirs',
 'themselves',
 'what',
 'which',
 'who',
 'whom',
 'this',
 'that',
 "that'll",
 'these',
 'those',
 'am',
 'is',
 'are',
 'was',
 'were',
 'be',
 'been',
 'being',
 'have',
 'has',
 'had',
 'having',
 'do',
 'does',
 'did',
 'doing',
 'a',
 'an',
 'the',
 'and',
 'but',
 'if',
 'or',
 'because',
 'as',
 'until',
 'while',
 'of',
 'at',
 'by',
 'for',
 'with',
 'about',
 'against',
 'between',
 'into',
 'through',
 'during',
 'before',
 'after',
 'above',
 'below',
 'to',
 'from',
 'up',
 'down',
 'in',
 'out',
```

#### Task 5
Remove the stops words in the lemmatized sentence 'lem_sent'.

Expected Output:

![title](img5.png)

<!-- grade_id: Task5, points: 0, grade: false, solution: true -->
```python
lem_sent_remove_stopwords = [word for word in lem_sent if word not in stopwords.words('english')]
# YOUR CODE HERE
# raise NotImplementedError()

lem_sent_remove_stopwords # Only important (not so common) words are left
```

Output:
```
['tokenization',
 'process',
 'segmenting',
 'running',
 'text',
 'sentence',
 'word',
 'essence',
 'task',
 'cutting',
 'text',
 'piece',
 'called',
 'token']
```

<!-- grade_id: Task5_1, points: 2, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
assert lem_sent_remove_stopwords == ['tokenization',
 'process',
 'segmenting',
 'running',
 'text',
 'sentence',
 'word',
 'essence',
 'task',
 'cutting',
 'text',
 'piece',
 'called',
 'token'], "Task 5 test(s) failed"


```

<!-- grade_id: cell-18f28aab13017ff9, points: 3, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

### Lessson 2 - Minimum Edit distance

#### Question 2.1
Implement the edit distance algorithm in Python

#### Task 6

Read through for understanding. 

```python
import numpy as np
def min_edit_distance(source, target, ins_cost = 1, del_cost = 1, rep_cost = 2):
    '''
    Input: 
        source: a string corresponding to the string you are starting with
        target: a string corresponding to the string you want to end with
        ins_cost: an integer setting the insert cost
        del_cost: an integer setting the delete cost
        rep_cost: an integer setting the replace cost
    Output:
        D: a matrix of len(source)+1 by len(target)+1 containing minimum edit distances
        med: the minimum edit distance (med) required to convert the source string to the target
    '''
    # use deletion and insert cost as  1
    m = len(source) 
    n = len(target) 
    #initialize cost matrix with zeros and dimensions (m+1,n+1) 
    D = np.zeros((m+1, n+1), dtype=int) 
    
    ### START CODE HERE (Replace instances of 'None' with your code) ###
    
    # Fill in column 0, from row 1 to row m, both inclusive
    for row in range(0,m+1): # Replace None with the proper range
        D[row,0] = row
        
    # Fill in row 0, for all columns from 1 to n, both inclusive
    for col in range(0,n+1): # Replace None with the proper range
        D[0,col] = col
        
    # Loop through row 1 to row m, both inclusive
    for row in range(1,m+1):
        
        # Loop through column 1 to column n, both inclusive
        for col in range(1,n+1):
            
            # Intialize r_cost to the 'replace' cost that is passed into this function
            r_cost = rep_cost
            
            # Check to see if source character at the previous row
            # matches the target character at the previous column, 
            if source[row-1] == target[col-1]: # Replace None with a proper comparison
                # Update the replacement cost to 0 if source and target are the same
                r_cost = 0
                
            # Update the cost at row, col based on previous entries in the cost matrix
            # Refer to the equation calculate for D[i,j] (the minimum of three calculated costs)
            D[row,col] = min([D[row-1, col-1]+r_cost, D[row,col-1]+ins_cost, D[row-1, col]+del_cost])
            
    # Set the minimum edit distance with the cost found at row m, column n 
    med = D[row,col]
    
    ### END CODE HERE ###
    return D, med
```

What is the minimum distance between 'dare' and 'dear'?

```python
print(min_edit_distance("dare", "dear", ins_cost = 1, del_cost = 1, rep_cost = 2)[1])
```

Output:
```
2
```

#### Task 7 - Build a Spell checker
- Download the files 'correct_words.txt' and 'incorrect_words.txt'
- As the name implies, 'correct_words' contain a list of correctly spelled words, while 'incorrect_words' is a list of incorrectly spelled words.
- Extract your vocabulary from 'correct_words.txt'. All words in the file are to be in the vocabulary.
- For each incorrect word in 'correct_words.txt', the task is to suggest a correct spelling from the list of correct words in 'correct_words.txt'.
   - Complete the function 'correct_spelling'. 
     - It takes in an incorrect word and the list of correct words (vocabulary)
     - It returns a new_word suggestion that is closest to the incorrect word from the vocabulary
        - Calculate the minimum edit distance between the incorrect word and all correct words, to find the closest word. Based on your calculation, suggest the possible intended word.
        - Use the min_edit_distance function given
     
A sample of the expected output is given:

![title](img6.png)

<!-- grade_id: Task7_a, points: 0, grade: false, solution: true -->
```python
with open('correct_words.txt', 'r') as file:
    #data = file.read().replace('\n'," ")
    data = file.readlines()
    vocabulary = [correct.replace("\n","") for correct in data]# complete the code
#     raise NotImplementedError()

with open('incorrect_words.txt', 'r') as file:
    #data = file.read().replace('\n'," ")
    data = file.readlines()
    incorrect_words = [a.replace('\n',"") for a in data ]
```

<!-- grade_id: Task7_a1, points: 2, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

<!-- grade_id: Task7b, points: 0, grade: false, solution: true -->
```python
def correct_spelling(word, vocabulary):
    # YOUR CODE HERE
    n = len(vocabulary)
    distances = [None]*n
    for i in range(n):
        correct_word = vocabulary[i]
        distances[i] = min_edit_distance(word,correct_word)[1]
    new_word = vocabulary[np.argmin(distances)]
    return new_word

```

<!-- grade_id: Task7b_1, points: 1, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
assert correct_spelling("abilitey",vocabulary) == "ability", "Task 7b test(s) failed"
print("Test(s) passed")
```

Output:
```
Test(s) passed
```

<!-- grade_id: Task7b_2, points: 1, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
assert correct_spelling("annonsment",vocabulary) == "announcement", "Task 7b test(s) failed"
print("Test(s) passed")
```

Output:
```
Test(s) passed
```

<!-- grade_id: Task7b_3, points: 1, grade: true, solution: false -->
```python
# RUN THIS CELL TO TEST YOUR CODE/FUNCTION
assert correct_spelling("cheet",vocabulary) == "cheat", "Task 7b test(s) failed"
print("Test(s) passed")
```

Output:
```
Test(s) passed
```

<!-- grade_id: Task7b_4, points: 1, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

<!-- grade_id: Task7b_5, points: 1, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

<!-- grade_id: Task7b_6, points: 1, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

<!-- grade_id: Task7b_7, points: 1, grade: true, solution: false -->
```python
#IGNORE THIS CELL, PERFORMING BACKGROUND TASKS
```

```python
for incorrect_word in incorrect_words:
    print(incorrect_word, correct_spelling(incorrect_word,vocabulary), " ")
```

Output:
```
abilitey ability  
abouy about  
absorbtion absorption  
accidently accidentally  
accomodate accommodate  
accosinly daily  
acommadate accommodate  
acord accord  
adultry adultery  
aggresive aggressive  
alchohol alcohol  
alchoholic alcoholic  
allieve alive  
amature amateur  
ambivilant ambivalent  
amification amplification  
amourfous amorphous  
annoint anoint  
annonsment announcement  
annuncio announce  
anonomy anatomy  
anotomy anatomy  
anynomous anonymous  
appreceiated appreciated  
appresteate appreciate  
aquantance acquaintance  
aratictature architecture  
archeype archetype  
aricticure arctic  
artic arctic  
asentote asymptote  
ast at  
asterick asterisk  
asymetric asymmetric  
atentively attentively  
autoamlly anatomy  
bankrot bankrupt  
basicly basically  
batallion battalion  
bbrose browse  
beauro bureau  
beaurocracy bureaucracy  
beggining beginning  
beging beginning  
behavior behavior  
behaviour behaviour  
beleive believe  
belive believe  
beni
```

#### Some thoughts
What do you think about this spell checker?
- The spell checker performs relatively well. Even though certain words suggested are incorrect, majority of the suggestions are correct and make sense.


Did you identify any suggestion that seems incorrect?
- Yes, some of these words are outlined in <b>Task 7b</b>. <br>

Can you identify any drawbacks? 
- A correct word that is not in the dictionary will never be suggested.
- We also iterate through the entire vocabulary and then find the minimum edit distance. This is inefficient.

For each drawback identified can you think of a solution?
- If we increase the size of the dictionary, more apt words that are not in the dictionary can be suggested.
- We can use BK-Trees, which will be eliminate the need to iterate through the entire vocabulary for each spell check.

##### Hints
Checking through an entire vocabulary is computationally expensive, perhaps it can be improved to only check a subset of words.
Also, no context of the word is factored in, what if we have the word 'iat', its edit distance with 'bat', 'cat', 'mat', and  'eat' are all the same. How can one make a choice?

#### Task 7b (Ungraded -   for learning purposes)

List some suggested words that you think are incorrect. Give the incorrect word, the suggested word based on the minimum edit distance and the most appropriate word from the vocabulary.

- Incorrect word: defenly   Correct word given: defence  Appropriate word: definitely
- Incorrect word: hifin   Correct word given: john  Appropriate word: hyphen
- Incorrect word: hifine   Correct word given: fiery  Appropriate word: hyphen
- Incorrect word: misilous   Correct word given: insiduous  Appropriate word: miscellaneous
- Incorrect word: rudemtry   Correct word given: destroy  Appropriate word: rudimentary
- Incorrect word: snuck   Correct word given: sonic  Appropriate word: sneaked
- Incorrect word: taht   Correct word given: at  Appropriate word: that
- Incorrect word: uncoscience   Correct word given: inconvenience  Appropriate word: unconscious
- Incorrect word: unformanlly   Correct word given: generally  Appropriate word: unfortunately

#### Task 8  ( (Ungraded -   for learning purposes)
Based on any of your identified drawbacks, create another spell checker that is an improvement on the one built in Task 7.
Is it able to give a better suggestion of any of the words identified in Task 7b?

- I could not think of a manner to improve this current spell checker. However, I did some research into BK-Trees. 
- In BK-Trees, we build a tree of our vocabulary and suggest words that have an edit distance less than or equal to a defined threshold. 
- However, this will return multiple possible replacements for a word and not a single correct word.
- The BK-Tree solves the problem of inefficeincy of the current spell checker , and cannot solve the problem of giving a more appropriate word, since it is also built using the edit distance like the spell checker.

### End of Lab